---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Before we can start to analyse a time series, we have to get one. 
A nice and easy way is to import the data directly form various internet sources such as 
Yahoo Finance or the Federal Reserve Bank of St. Louis (who provide plenty of economic datasets) using the `quantmod` package. 


`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
